# DATASCI 507, Week 1 practice

### Problem 1
Write a Python function that prints "Hello, World!".

```python
def print_hello_world():
    # YOUR CODE HERE
    raise NotImplementedError()

print_hello_world()
```

Output:
```
Cell In[1], line 5
      1 def print_hello_world():
      2     # YOUR CODE HERE
      3     raise NotImplementedError()
      4 
----> 5 print_hello_world()

Cell In[1], line 3, in print_hello_world()
      1 def print_hello_world():
      2     # YOUR CODE HERE
----> 3     raise NotImplementedError()

NotImplementedError:
```

```python
print("Hello World")
```

### Problem 2
Write a function that takes two integers as argumens. Return the first argument if it is a positive number. Otherwise, return the sum of both arguments.

<!-- grade_id: cell-db7233a2d95c6406, points: 0, grade: false, solution: true -->
```python
def postive_or_sum(x, y):
    # YOUR CODE HERE
    raise NotImplementedError()

assert postive_or_sum(1,2) == 1
```

<!-- grade_id: cell-d525ecbed5402db1, points: 1, grade: true, solution: false -->
```python
assert postive_or_sum(42, 16) == 42
assert postive_or_sum(33, 13) == 33
assert postive_or_sum(-33, 13) == -20
```

```python
assert postive_or_sum(3,4) == 7
```